### Exercise 4: NMIST classification

Name and student id: 

```python
import numpy as np
from numpy import genfromtxt
import matplotlib.pyplot as plt
from sklearn.utils.extmath import cartesian
from sklearn.datasets import fetch_openml
from sklearn import preprocessing
from sklearn.preprocessing import StandardScaler
from sklearn.metrics import confusion_matrix
from sklearn.model_selection import train_test_split
from sklearn.svm import SVC
from sklearn import tree
from sklearn.ensemble import RandomForestClassifier
```

```python
x, y = fetch_openml('mnist_784', version=1, return_X_y=True) # download takes a while 
```

Standardize feature data and split the data into training and test sets with 30% for training.

<!-- grade_id: cell-209025c51cb01994, points: 0, grade: false, solution: true -->
```python
# scaled_x = ...
# x_train, x_test, y_train, y_test = ...
# YOUR CODE HERE
raise NotImplementedError()
print(x_train.shape)
```

<!-- grade_id: cell-12fbc75302e1add3, points: 2, grade: true, solution: false -->
```python
assert(x_train.shape == (49000, 784))
```

Train a decision tree classifier model (assign to variable model) with default parameters. Compute training and test scores and assign them to training_score and test_score. Also print the confusion matrix for both training and test sets.

<!-- grade_id: cell-c003e2eba3246ea7, points: 0, grade: false, solution: true -->
```python
# model = ...
# training_score = ...
# test_score = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4dd090649c9d931d, points: 7, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(training_score, 1.0, decimal=1)
np.testing.assert_almost_equal(test_score, 0.87, decimal=1)
```

How many times in validation set
- digit 3 gets misclassified?
- digit 4 gets misclassified as 7?

What is the most misclassified digit?

Please answer in cell below.

<!-- grade_id: cell-55c2755f23c8ffd9, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

Train a random forest classifier model (assigned to variable model) with default parameters. Compute training and test scores and assign them to training_score and test_score. Also print the confusion matrix for both training and test sets.

<!-- grade_id: cell-79b9eddba0cbcb28, points: 0, grade: false, solution: true -->
```python
# model = ...
# training_score = ...
# test_score = ...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0d5a87fccfa891dc, points: 7, grade: true, solution: false -->
```python
np.testing.assert_almost_equal(training_score, 0.999, decimal=2)
np.testing.assert_almost_equal(test_score, 0.944, decimal=2)
```